## Problem: Sliding window

Implement a sliding window for an arbitrary input list. The function should take
the window width and the window increment as inputs and should produce a
sequence of overlapping lists from the input list. For example, given
`x=range(15)`, the following is the output given a window width of `5` and
window increment of `2`.

    [[0, 1, 2, 3, 4],
     [2, 3, 4, 5, 6],
     [4, 5, 6, 7, 8],
     [6, 7, 8, 9, 10],
     [8, 9, 10, 11, 12],
     [10, 11, 12, 13, 14]]

In the event that the input parameters do not yield a complete set of even
sublists, just truncate the ragged tail. For example,

    >>> slide_window(range(18),5,2)
    [[0, 1, 2, 3, 4],
     [2, 3, 4, 5, 6],
     [4, 5, 6, 7, 8],
     [6, 7, 8, 9, 10],
     [8, 9, 10, 11, 12],
     [10, 11, 12, 13, 14],
     [12, 13, 14, 15, 16]]

<!-- grade_id: cell-feb98530bbfed6a6, points: 0, grade: false, solution: true -->
```python
def slide_window(x,width,increment):
    '''
    take the window width and the window increment as inputs and should produce a sequence of overlapping lists from the input list
    Args:
         x, width, increment
    type:
        range or list,int,int
    output:
        produce a sliding window for given parameters
    type: 
        list
    '''
    assert isinstance(x,list) or isinstance(x,range),'x is not a list'
    assert isinstance(width,int),'width is not a integer'
    assert isinstance(increment,int),'increment is not a integer'
    assert width > increment,'width less than increment'
    assert increment > 0 ,'increment should be positive'
    assert width >= 1,'invalid width'
    assert len(x) >= width,'width larger than length of input list'
    w = []
    for i in x:
        w.append(i)
    slideWindow = []
    intial_num = 0
    current_num = width
    while (intial_num + width) <= len(w):
        tempList = []
        for i in range(intial_num,current_num): 
            tempList.append(i)
            # print(tempList)
        slideWindow.append(tempList)
        intial_num = intial_num + increment
        current_num = current_num + increment
    

    return slideWindow




```

<!-- grade_id: cell-ca9bd3475cce33f2, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-21ef5921c0f7a2d6, points: 2, grade: true, solution: false -->
```python
assert slide_window.__doc__ is not None
```

```python

```