Recall the data from last time

    data=[5,4,6,1,9,0,3,9,2,7,10,8,4,7,1,2,7,6,5,2,8,2,0,1,1,1,2,10,6,2]

Write three columns to a comma-separated file:

    data_value, data_value**2, (data_value+data_value**2)/3.

<!-- grade_id: cell-143b614df20a8f3b, points: 0, grade: false, solution: true -->
```python
def write_columns(data,fname):
    '''
    Given data as a list, write three columns to fname.
    Args:
        data, fname
    type:
        list, str
    '''
    assert isinstance(data,list),'data is not a list'
    assert isinstance(fname,str),'data is not a string'
    for element in data:
        isinstance(element,int) or isinstance(element,float),'element in data is not integer or float'
    with open(fname,'w') as inputFile:
        for data_value in data:
            print(data_value,',', data_value**2,',', (data_value+data_value**2)/3,file = inputFile)
    
    



```

<!-- grade_id: cell-17cf039687237cdd, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
import numpy as np
from nose.tools import assert_almost_equal
data=[5,4,6,1,9,0,3,9,2,7,10,8,4,7,1,2,7,6,5,2,8,2,0,1,1,1,2,10,6,2]
fname = 'test_data.csv'
write_columns(data,fname)
assert_almost_equal(np.loadtxt(fname,delimiter=',').sum(),1328,2)
    
### END  TESTS
```

```python

```

```python

```

```python

```

```python

```