Download this [corpus of 10,000 common English words](http://localhost:8888/?token=a6840a33f9885242e17dacd3d5ed7114465feb61d54ced90) and do the following:


- What is the longest word?
- What is the longest word that starts with s
- What is the most common starting letter?
- What is the most common ending letter?

<!-- grade_id: cell-9b9db364a733cab0, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e9f6429b2bc8dfbc, points: 0, grade: false, solution: true -->
```python
# you can use this bit of code to download the words from the corpus
from urllib.request import urlopen
import re
from collections import Counter

u='https://storage.googleapis.com/class-notes-181217.appspot.com/google-10000-english-no-swears.txt'
response = urlopen(u)
words = [i.strip().decode('utf8') for i in response.readlines()]

def get_longest_word(words):
    '''
    get the longest word for a given input file
    Args:
        input words list
    type: 
        str
    Return: 
        list with longest word(s)
    type: 
        list
    '''
    assert isinstance(words,list),'words is not a list'
    for element in words:
        assert isinstance(element,str),'the element in the list is not a string'
    max_len = len(max(words, key=len))
    finalList = [word for word in words if len(word) == max_len]
    return finalList[0]





def get_longest_words_startswith(words,starts):
    '''
    get the longest words starts with
    Args:
        input words list, starts letter
    type: 
        list, string
    Return:
        the longest word with specific starting letter
    type: 
        string
    '''
    assert isinstance(words,list),'words is not a list'
    assert isinstance(starts,str),'starting letter is not a string'
    for element in words:
        assert isinstance(element,str),'the element in the list is not a string'
    wordsListWithStarts = [word for word in words if word[0] == 's']
    for word in wordsListWithStarts:
        max_len = len(max(wordsListWithStarts,key=len))
    finalList = [word for word in wordsListWithStarts if len(word) == max_len]
    return finalList[0]
    

    
def get_most_common_start(words):
    '''
    get the most common letter start with
    Args:
        input words list
    type:
        list
    Return:
        the most common letter start with
    type: 
        string
    '''
    assert isinstance(words,list),'words is not a list'
    for element in words:
        assert isinstance(element,str),'the element in the list is not a string'
    firstLetter = [item[0] for item in words]
    word_counts = Counter(firstLetter)
    mostStartingLetter = word_counts.most_common(1)
    return mostStartingLetter[:1][0][0] #return a single character...
        



def get_most_common_end(words):
    '''
    get the most common end letter for the given words list
    Args:
        input words list
    type:
        list
    Return:
        the most common end letter
    type:
        string
    '''
    assert isinstance(words,list),'wors is not a list'
    for element in words:
        isinstance(element,str),'the element in the list is not a string'
    list1 = []
    list2 = []
    list3 = []
    a = 0
    b = 0
    for i in range(len(words)):
        if words[i][len(words[i])-1] not in list1:
            list1.append(words[i][len(words[i])-1])
    list2 = sorted(list1)
    for j in range(len(list2)):
        count=0
        for i in range(len(words)):
            if words[i][len(words[i])-1]==list2[j]:
                count+=1
        list3.append(count)
    for j in range(len(list3)):
        if a <= list3[j]:
            a = list3[j]
            b = j
    return list2[b]


```

<!-- grade_id: cell-7559f8c989e26885, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert get_longest_words_startswith(words,'s') in  ['sustainability','simultaneously', 
'specifications', 
'sustainability', 
'superintendent' ]
### END  TESTS
```

<!-- grade_id: cell-3cf6876fc57af686, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert get_most_common_start(words)=='c'
### END  TESTS
```

<!-- grade_id: cell-f1359f41f9db52bd, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert get_most_common_end(words)=='s'
### END  TESTS
```

```python

```

```python

```

```python

```